<problem display_name="Knowledge Check 3" markdown="null" max_attempts="2">
  <script type="loncapa/python"> 
import re 
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 1
</script>
  <span id="anchor_change_my_id_for_each_one_on_the_page"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_change_my_id_for_each_one_on_the_page") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>What are some of the different ways that a building may fail?</p>
  <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>Here are the main ways that a building may fail: bending (too much bending due to wind or earthquake), buckling (too much vertical load), shear (short beams and buildings will be more susceptible to shear failures), overturning, and sliding (overturning and sliding are mainly a function of connections at the base not being strong enough).</p>
    </div>
  </solution>
</problem>
