<problem display_name="What do you think of Owl's Solution?" markdown="null" max_attempts="3">
  <script type="loncapa/python"> 
import re 
   
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 0

   
</script>
  <span id="anchor_change_my_id_for_each_one_on_the_page"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_change_my_id_for_each_one_on_the_page") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>
    <strong><font color="black">What do you think of Owl's design? What would you have designed?</font></strong>
  </p>
   <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>Owl designed a propped cantilever beam for the diving board. Other types of beam Owl might have designed include a simply-supported beam and a cantilever beam. A simply-supported beam, maybe spanning between the two trees would have had less flexibility and would have been more difficult to jump from. A cantilever beam would have worked but would require a strong, moment-resisting connection, which tend to be more difficult to build and may have damaged the tree.</p>
      <p/>
    </div>
  </solution>
</problem>
