<problem display_name="Knowledge Check 5" markdown="null" max_attempts="2">
  <script type="loncapa/python"> 
import re 
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 1
</script>
  <span id="anchor_KC5"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_KC5") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>Why may an engineer create a simplified model of a building? Isn’t a more detailed model better?</p>
  <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>Simplified models can tell you a lot about the behavior more quickly so they are useful as preliminary design tools. Detailed models are great too but it is easy to get lost in the large quantity of numbers associated with these models.</p>
      <p> </p>
    </div>
  </solution>
</problem>
