<problem display_name="Knowledge Check 12" markdown="null" max_attempts="2">
  <script type="loncapa/python"> 
import re 
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 1
</script>
  <span id="anchor_KC12"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_KC12") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>During the Mexico City earthquake of 1985 certain buildings collapsed, while others remained standing. What was the issue with the buildings that collapsed? How could these collapses have been avoided?</p>
  <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>The 1985 Mexico City earthquake had a predominant frequency that matched the natural frequency of many of the buildings in the area in the 5-15 story range. Shorter and taller buildings were not as affected and remained standing. After this earthquake the codes changed and engineers started paying attention to the frequencies of both buildings and earthquakes.</p>
      <p> </p>
    </div>
  </solution>
</problem>
