<problem display_name="Knowledge Check 19" markdown="null" max_attempts="2">
  <script type="loncapa/python"> 
import re 
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 1
</script>
  <span id="anchor_KC19"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_KC19") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>What are some of the differences involved with designing buildings (beams, columns, and walls) using glass rather than a more traditional building material?</p>
  <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
        <p>Glass is a brittle and unpredictable material so there are many precautions that must be taken when designing with it such as tempering the glass and laminating multiple layers together. The connections between pieces of glass are critical.</p>
      <p/>
    </div>
  </solution>
</problem>
