<problem display_name="Knowledge Check 10" markdown="null" max_attempts="2">
  <script type="loncapa/python"> 
import re 
def checkAnswer4(expect, ans): 
  response = re.search('', ans) 
  if response: 
    return 1 
  else: 
    return 1
</script>
  <span id="anchor_KC10"/>
  <script type="text/javascript"> 
    /* The object here is to replace the single line input with a textarea */ 
   (function() { 
    var elem = $("#anchor_KC10") 
        .closest("div.problem") 
        .find(":text"); 
    /* There's CSS in the LMS that controls the height, so we have to override here */ 
    var textarea = $('<textarea style="height:150px" rows="20" cols="70"/>'); 
    console.log(elem); 
    console.log(textarea); 
    //This is just a way to do an iterator in JS 
    for (attrib in {'id':null, 'name':null}) { 
        textarea.attr(attrib, elem.attr(attrib)); 
    } 
    /* copy over the submitted value */ 
    textarea.val(elem.val()) 
    elem.replaceWith(textarea); 
     
   })(); 
  </script>
  <p>Is it better to design a building that is very stiff, meaning it deflects or moves very little, or one that is more flexible, meaning it deflects or moves quite a bit?</p>
  <customresponse cfn="checkAnswer4">
    <textline size="40" correct_answer="" label="Problem Text"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>It depends! For some sites a stiffer building may be better while at others a more flexible building may be better. It depends on the expected earthquake motions. The goal is to tune the stiffness and mass of the building such that they are different from the expected input motions.</p>
      <p> </p>
    </div>
  </solution>
</problem>
